```python

```

<img align="center" src="images/course.png" width="800">

# 16720 (B)  3D Reconstruction - Assignment 5 - q2
    Instructor: Kris                          TAs: Arka, Jinkun, Rawal, Rohan, Sheng-Yu

```python
# Helper functions for this assignment. DO NOT MODIFY!!!
"""
Helper functions.
"""

import numpy as np
import matplotlib.pyplot as plt
import matplotlib
import scipy
import nbimporter
from q1 import eightpoint, camera2, _epipoles, toHomogenous

def epipolarMatchGUI(I1, I2, F):
    matplotlib.use('TkAgg')
    e1, e2 = _epipoles(F)

    sy, sx, _ = I2.shape

    f, [ax1, ax2] = plt.subplots(1, 2, figsize=(12, 9))
    ax1.imshow(I1)
    ax1.set_title('Select a point in this image')
    ax1.set_axis_off()
    ax2.imshow(I2)
    ax2.set_title('Verify that the corresponding point \n is on the epipolar line in this image')
    ax2.set_axis_off()

    while True:
        plt.sca(ax1)
        x, y = plt.ginput(1, mouse_stop=2)[0]

        xc = int(x)
        yc = int(y)
        v = np.array([xc, yc, 1])
        l = F.dot(v)
        s = np.sqrt(l[0]**2+l[1]**2)

        if s==0:
            error('Zero line vector in displayEpipolar')

        l = l/s;

        if l[0] != 0:
            ye = sy-1
            ys = 0
            xe = -(l[1] * ye + l[2])/l[0]
            xs = -(l[1] * ys + l[2])/l[0]
        else:
            xe = sx-1
            xs = 0
            ye = -(l[0] * xe + l[2])/l[1]
            ys = -(l[0] * xs + l[2])/l[1]

        # plt.plot(x,y, '*', 'MarkerSize', 6, 'LineWidth', 2);
        ax1.plot(x, y, '*', MarkerSize=6, linewidth=2)
        ax2.plot([xs, xe], [ys, ye], linewidth=2)

        # draw points
        x2, y2 = epipolarCorrespondence(I1, I2, F, xc, yc)
        ax2.plot(x2, y2, 'ro', MarkerSize=8, linewidth=2)
        plt.draw()

def plot_3D(P):
    matplotlib.use('TkAgg')
    fig = plt.figure()
    ax = fig.add_subplot(111, projection='3d')
    ax.scatter(P[:,0], P[:,1], P[:,2])
    while True:
        x, y = plt.ginput(1, mouse_stop=2)[0]
        plt.draw()

        

```

## Q2: Metric Reconstruction
You will compute the camera matrices and triangulate the 2D points to obtain the 3D scene structure. To obtain the Euclidean scene structure, first convert the fundamental matrix $\textbf{F}$ to an essential matrix $\textbf{E}$. Examine the lecture notes and the textbook to find out how to do this when the internal camera calibration matrices $\textbf{K}_1$ and $\textbf{K}_2$ are known; these are provided in `data/intrinsics.npz`.

### Q2.1: Essential Matrix (5 pt implementation)
Write a function to compute the essential matrix $\textbf{E}$ given $\textbf{F}$, $\textbf{K}_1$ and $\textbf{K}_2$ with the signature:

```
E = essentialMatrix(F, K1, K2)
```



<!-- grade_id: cell-75539a0c38616db4, points: 0, grade: false, solution: true -->
```python
def essentialMatrix(F, K1, K2):
    '''
    Q1.1: Compute the essential matrix E given the fundamental matrix and camera intrinsics
        Input:  F, fundamental matrix
                K1, internal camera calibration matrix of camera 1
                K2, internal camera calibration matrix of camera 2
        Output: E, the essential matrix
    '''
    # ----- TODO -----
    # YOUR CODE HERE
    # from page 228 Forsyth and Ponce
    E = K2.T @ F @ K1
    # raise NotImplementedError()
    return E


```

```python
# Printing out the essential matrix
np.set_printoptions(precision=4, suppress=1)
correspondence = np.load('data/some_corresp.npz') # Loading correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1 = plt.imread('data/im1.png')
im2 = plt.imread('data/im2.png')

F = eightpoint(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))
E = essentialMatrix(F, K1, K2)
print(E)
```

Output:
```
Optimization terminated successfully.
         Current function value: 0.000107
         Iterations: 8
         Function evaluations: 894
[[  -0.5069   68.6543 -371.9615]
 [  29.7107   -1.5472    9.6823]
 [ 372.9911    2.9855    0.1504]]
```

<!-- grade_id: q2_1_a, points: 5, grade: true, solution: false -->
```python
# Simple Tests to verify your implmentation:
correspondence = np.load('data/some_corresp.npz') # Loading correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
F = eightpoint(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))
E = essentialMatrix(F, K1, K2)
assert(np.linalg.matrix_rank(E) == 2)

# Hidden Tests
```

Output:
```
Optimization terminated successfully.
         Current function value: 0.000107
         Iterations: 8
         Function evaluations: 894
```

### Q2.2 Triangulation and find M2 (5 pt writeup, 20 pt implementation)
Given an essential matrix, it is possible to retrieve the projective camera matrices $\textbf{M}_1$ and $\textbf{M}_2$ from it.  Assuming $\textbf{M}_1$ is fixed at $[\textbf{I},0]$, $\textbf{M}_2$ can be retrieved up to a scale and four-fold rotation ambiguity. For details on recovering $\textbf{M}_2$, see section 7.2 in Szeliski. We have provided you with the function `camera2` to recover as the four possible $\textbf{M}_2$ matrices given $\textbf{E}$.

**Note:** The $\textbf{M}_1$ and $\textbf{M}_2$ here are projection matrices of the form:
$\textbf{M}_1 = \begin{bmatrix}
\textbf{I} | 0
\end{bmatrix} $ and $\textbf{M}_2 = \begin{bmatrix}
\textbf{R} | \textbf{t}
\end{bmatrix} $.

Using the above, write a function to triangulate a set of 2D coordinates in the image to a set of 3D
points with the signature:
```
    [w, err] = triangulate(C1, pts1, C2, pts2)
```
where `pts1` and `pts2` are the $N \times 2$ matrices with the 2D image coordinates and `w` is an $N \times 3$ matrix with the corresponding 3D points per row.  `C1` and `C2` are the $3 \times 4$ camera matrices. Remember that you will need to multiply the given intrinsics matrices with your solution for the canonical camera matrices to obtain the final camera matrices. Various methods exist for triangulation - probably the most familiar for you is based on least squares (see Szeliski Chapter 7 if you want to learn about other methods):

For each point $i$, we want to solve for 3D coordinates $\textbf{w}_i = \begin{bmatrix}x_i, y_i, z_i\end{bmatrix}^T$, such that when they are projected back to the two images, they are close to the original 2D points. To project the 3D coordinates back to 2D images, we first write $\textbf{w}_i$ in homogeneous coordinates, and compute $\mathbf{C}_1 \tilde{\textbf{w}_i}$ and $\mathbf{C}_2 \tilde{\textbf{w}_i}$ to obtain the 2D homogeneous coordinates projected to camera $1$ and camera $2$, respectively.

For each point $i$, we can write this problem in the following form:
\begin{align}
\mathbf{A}_i w_i = 0
\end{align}
where $\mathbf{A}_i$ is a $4\times 4$ matrix, and $\tilde{\textbf{w}_i}$ is a $4\times 1$ vector of the 3D coordinates in the homogeneous form. Then, you can obtain the homogeneous least-squares solution (discussed in class) to solve for each $\textbf{w}_i$.

Once you have implemented triangulation, you can check the performance by looking at the reprojection error:  $$ \texttt{err} = \sum_i ||\textbf{x}_{1i} - \widehat{\textbf{x}_{1i}}||^2 + ||\textbf{x}_{2i} - \widehat{\textbf{x}_{2i}}||^2$$
where $\widehat{\textbf{x}_{1i}} = Proj(\mathbf{C}_1, \textbf{w}_i)$ and $\widehat{\textbf{x}_{2i}} = Proj(\mathbf{C}_2, \textbf{w}_i)$.


**Note:** `C1` and `C2` here are projection matrices of the form:
$\mathbf{C}_1 = \mathbf{K}_1\mathbf{M}_1 = \mathbf{K}_1 \begin{bmatrix}
\textbf{I} | 0
\end{bmatrix} $ and $\mathbf{C}_2 =  \mathbf{K}_2\mathbf{M}_2 = \mathbf{K}_2 \begin{bmatrix}
\textbf{R} | \textbf{t}
\end{bmatrix}$.

Write a function, 
```
    [M2, C2, P] = findM2()
```
to load the correspondences and use triangulation to find the M2, C2, and P that minimize the reprojection error. 

<span style='color:red'>**Output:**</span> In your write-up: Write down the expression for the matrix $\mathbf{A}_i$

<!-- grade_id: cell-af5ab5f8a3142145, points: 0, grade: false, solution: true -->
```python
def triangulate(C1, pts1, C2, pts2):
    '''
    Q2.2: Triangulate a set of 2D coordinates in the image to a set of 3D points.
        Input:  C1, the 3x4 camera matrix
                pts1, the Nx2 matrix with the 2D image coordinates per row
                C2, the 3x4 camera matrix
                pts2, the Nx2 matrix with the 2D image coordinates per row
        Output: P, the Nx3 matrix with the corresponding 3D points per row
                err, the reprojection error.
    
    ***
    Hints:
    (1) For every input point, form A using the corresponding points from pts1 & pts2 and C1 & C2
    (2) Solve for the least square solution using np.linalg.svd
    (3) Calculate the reprojection error using the calculated 3D points and C1 & C2 (do not forget to convert from 
        homogeneous coordinates to non-homogeneous ones)
    (4) Keep track of the 3D points and projection error, and continue to next point 
    (5) You do not need to follow the exact procedure above. 
    '''
    # ----- TODO -----
    # YOUR CODE HERE
    N = pts1.shape[0]
    C11, C12, C13 = C1[0,:], C1[1,:], C1[2,:]
    C21, C22, C23 = C2[0,:], C2[1,:], C2[2,:] 
    err = 0
    P = []
    for i in range(N):
        (x, y) = pts1[i,:]
        (x_, y_) = pts2[i,:]
        A = np.array([y*(C13) - C12,
                      C11 - x*(C13),
                      y_*(C23) - C22,
                      C21 - x_*(C23)])
    
        assert A.shape==(4,4), 'A dims wrong'
        U, S, VT = np.linalg.svd(A)
        p = VT[-1,:] # last row of VT => last column of V (4,)
        p = p/p[-1] # [X,Y,Z,1] (wi in homogeneous coordinates)
        wi = p[:-1] # [X,Y,Z]
        P.append(wi)
        
        X1i_cap = C1@(p.T)
        X1i_cap = X1i_cap/X1i_cap[-1]
        X1i_cap = X1i_cap[:-1]

        X2i_cap = C2@(p.T)
        X2i_cap = X2i_cap/X2i_cap[-1]
        X2i_cap = X2i_cap[:-1]
        # calculating reprojection error
        err += np.linalg.norm(np.array([x,y]).T-X1i_cap)**2 + np.linalg.norm(np.array([x_,y_]).T-X2i_cap)**2

    print(A)
        
    P = np.array(P)
    # raise NotImplementedError()

    return P, err
    

def find_M2(F, pts1, pts2, intrinsics):
    '''
    Q2.2: Function to find the camera2's projective matrix given correspondences
        Input:  F, the pre-computed fundamental matrix
                pts1, the Nx2 matrix with the 2D image coordinates per row
                pts2, the Nx2 matrix with the 2D image coordinates per row
                intrinsics, the intrinsics of the cameras, load from the .npz file
        Output: [M2, C2, P] the computed M2 (3x4) camera projective matrix, C2 (3x4) K2 * M2, and the 3D points P (Nx3)
    
    ***
    Hints:
    (1) Loop through the 'M2s' and use triangulate to calculate the 3D points and projection error. Keep track 
        of the projection error through best_error and retain the best one. 
    (2) Remember to take a look at camera2 to see how to correctly reterive the M2 matrix from 'M2s'. 

    '''
    K1 = intrinsics["K1"]
    K2 = intrinsics["K2"]
    E = essentialMatrix(F, K1, K2)
    M1 = np.hstack((np.identity(3), np.zeros(3)[:,np.newaxis]))
    C1 = K1.dot(M1)
    M2s = camera2(E) # 4 possible M2 matrices (3, 4, 4)
    best_error = np.finfo('float').max # largest representable floating point number
    M2 = np.zeros_like(M2s[:,:,0])
    C2 = np.zeros_like(C1)
    # ----- TODO -----
    # YOUR CODE HERE
    n_M2s = M2s.shape[2]
    P_best = None
    for i in range(n_M2s):
        M2i = M2s[:,:,i]
        C2i = K2.dot(M2i)
        P, err = triangulate(C1, pts1, C2i, pts2)
        # print(err)
        if err<best_error and not np.any(P[:,2]<0, axis=0):
            best_error = err
            M2 = M2i
            C2 = C2i
            P_best = P

    P = P_best
    # raise NotImplementedError()
    print(f"Best Error {best_error}")
    return M2, C2, P
    
```

```python
# Running the find_M2 funciton:
np.set_printoptions(precision=4, suppress=1)
correspondence = np.load('data/some_corresp.npz') # Loading correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1, im2 = plt.imread('data/im1.png'), plt.imread('data/im2.png')

F = eightpoint(pts1, pts2, M = np.max([*im1.shape, *im2.shape]))
M2, C2, P = find_M2(F, pts1, pts2, intrinsics)
print(f"M2: {M2}")
print(f"C2 {C2}")
```

Output:
```
Optimization terminated successfully.
         Current function value: 0.000107
         Iterations: 8
         Function evaluations: 894
[[    0.     -1525.9      -87.87       0.    ]
 [ 1520.4        0.        66.32       0.    ]
 [  -27.1612 -1527.0144   -65.5751  1518.5853]
 [-1520.2631    33.8465   -62.757    -34.0881]]
[[    0.     -1525.9      -87.87       0.    ]
 [ 1520.4        0.        66.32       0.    ]
 [  -27.1612 -1527.0144   -65.5751 -1518.5853]
 [-1520.2631    33.8465   -62.757     34.0881]]
[[    0.     -1525.9      -87.87       0.    ]
 [ 1520.4        0.        66.32       0.    ]
 [   51.2048 -1449.186   -483.7859  1518.5853]
 [ 1519.7233    32.9612    75.0886   -34.0881]]
[[    0.     -1525.9      -87.87       0.    ]
 [ 1520.4        0.        66.32       0.    ]
 [   51.2048 -1449.186   -483.7859 -1518.5853]
 [ 1519.7233    32.9612    75.0886    34.0881]]
Best Error 351.8979672663218
M2: [[ 0.9994  0.0333  0.006  -0.026 ]
 [-0.0337  0.9653  0.2589 -1.    ]
 [ 
```

<!-- grade_id: q2_2_a, points: 5, grade: true, solution: false -->
```python
# Simple Tests to verify your implmentation:
M1 = np.hstack((np.identity(3), np.zeros(3)[:,np.newaxis]))
C1 = K1.dot(M1)
C2 = K2.dot(M2)
P_test, err = triangulate(C1, pts1, C2, pts2)
print('error:', err)
try:
    assert(err < 500)
    print('Test passed!')
except:
    raise AssertionError('Test Failed: error greater than 500!')
```

Output:
```
[[    0.     -1525.9      -87.87       0.    ]
 [ 1520.4        0.        66.32       0.    ]
 [   51.2048 -1449.186   -483.7859  1518.5853]
 [ 1519.7233    32.9612    75.0886   -34.0881]]
error: 351.8979672663218
Test passed!
```

<!-- grade_id: q2_2_b, points: 15, grade: true, solution: false -->
```python
# Hidden Tests
```

### Q2.3 Epipolar Correspondence (5 pt writeup, 10 pt implementation)

You will now create a 3D visualization of the temple images. By treating our
two images as a stereo-pair, we can triangulate corresponding points in each
image, and render their 3D locations.



Implement a function with the
signature:
```
    [x2, y2] = epipolarCorrespondence(im1, im2, F, x1, y1)
```

This function takes in the $x$ and $y$ coordinates of a pixel on `im1` and your fundamental matrix $\textbf{F}$, and returns the coordinates of the pixel on `im2` which correspond to the input point. The match is obtained by computing the similarity of a small window around the $(x_1, y_1)$ coordinates in `im1` to various windows around possible matches in the `im2` and returning the closest.

Instead of searching for the matching point at every possible location in `im2`, we can use $\textbf{F}$ and simply search over the set of pixels that lie along the epipolar line (recall that the epipolar line passes through a single
point in `im2`  which corresponds to the point $(x_1, y_1)$ in `im1`!.

There are various possible ways to compute the window similarity. For this assignment, simple methods such as the Euclidean or Manhattan distances between the intensity of the pixels should suffice.  See Szeliski Chapter 11, on stereo matching, for a brief overview of these and other methods.


**Implementation hints:**
- Experiments with various window sizes.
- It may help to use a Gaussian weighting of the window, so that the center
  has greater influence than the periphery.
- Since the two images only differ by a small amount, it might be beneficial to consider matches for which the distance from $(x_1, y_1)$ to $(x_2, y_2)$ is small.

To help you test your `epipolarCorrespondence`, we have included a helper function `epipolarMatchGUI`, which takes in two images the fundamental matrix.

This GUI allows you to click on a point in `im1`, and will use your function to display the corresponding point in `im2`. See:

<img align="center" src="images/q3gui_disp_arun.png" width="800">

It's not necessary for your matcher to get **every** possible point right, but it should get easy points (such as those with distinctive corner-like windows).

It should also be good enough to render an intelligible representation in the next question.

<span style='color:red'>**Output:**</span> In your write-up, include a screenshot of `epipolarMatchGUI`
with some detected correspondences.


<!-- grade_id: cell-5675d5028163b272, points: 0, grade: false, solution: true -->
```python
import scipy.ndimage
from scipy.ndimage.filters import gaussian_filter


def epipolarCorrespondence(im1, im2, F, x1, y1):
    '''
    Q2.3 3D visualization of the temple images.
        Input:  im1, the first image
                im2, the second image
                F, the fundamental matrix
                x1, x-coordinates of a pixel on im1
                y1, y-coordinates of a pixel on im1
        Output: x2, x-coordinates of the pixel on im2
                y2, y-coordinates of the pixel on im2
    
    ***
    Hints:
    (1) Given input [x1, x2], use the fundamental matrix to recover the corresponding epipolar line on image2
    (2) Search along this line to check nearby pixel intensity (you can define a search window) to 
        find the best matches
    (3) Use guassian weighting to weight the pixel simlairty
    
    '''
    # ----- TODO -----
    # YOUR CODE HERE
    e1, e2 = _epipoles(F)
    sy, sx, _ = im2.shape

    xc = x1
    yc = y1
    v = np.array([xc, yc, 1])
    l = F.dot(v)
    s = np.sqrt(l[0]**2+l[1]**2)

    if s==0:
        error('Zero line vector in displayEpipolar')

    l = l/s

    if l[0] != 0:
        ye = sy-1
        ys = 0
        xe = -(l[1] * ye + l[2])/l[0]
        xs = -(l[1] * ys + l[2])/l[0]
    else:
        xe = sx-1
        xs = 0
        ye = -(l[0] * xe + l[2])/l[1]
        ys = -(l[0] * xs + l[2])/l[1]

    # points_along_el = []
    # for x in range(np.ceil(xs).astype(int), np.floor(xe).astype(int)):
    #     y = -(l[0] * x + l[2])/l[1]
    #     y = int(y)
    #     pt = [x,y]
    #     points_along_el.append(pt)
    N = max( (ye-ys), (xe-xs) )
    x2_list = np.linspace(xs, xe, N)
    y2_list = np.linspace(ys, ye, N)

    # x2_list = np.rint(x2_list).astype(int)
    # y2_list = np.rint(y2_list).astype(int)
    x2_list = np.array(x2_list).astype(int)
    y2_list = np.array(y2_list).astype(int)

    
    k = 51 # window size (k,k)
    sigma = 31 # std dev in gaussian weighting
    assert k%2 !=0, 'window size should be odd'
    assert x1 >= k//2 and y1 >= k//2 and x1 + k//2 <= sx-1 and y1 + k//2 <= sy-1, 'point too close to corner or is out of bounds'

    patch_1 = im1[y1 - k//2: y1 - k//2 + k, x1 - k//2: x1 - k//2 + k, :]
    patch_1 = np.asarray(patch_1)

    # gauss_filter = create_gauss_filter(k, sigma)
    # kernel = makeGaussianFiler(k, sigma)
    # kernel /= np.sum(kernel)
    # kernel = np.asarray(kernel)
    # kernel = np.dstack( ( kernel, kernel, kernel  )  )
    dist_min = np.inf
    best_pt = None
    # for pt in points_along_el:
    for pt in list(zip(x2_list, y2_list)):

        x2,y2 = pt
        if x2 >= k//2 and y2 >= k//2 and x2 + k//2 <= sx-1 and y2 + k//2 <= sy-1:
            patch_2 = im2[y2 - k//2: y2 - k//2 + k, x2 - k//2: x2 - k//2 + k, :]
            patch_2 = np.asarray(patch_2)
            diff = patch_1-patch_2
            dist = np.linalg.norm(diff)

            
            if dist < dist_min:
                best_pt = pt
                dist_min = dist

    x2,y2 = best_pt

    return x2, y2

```

```python
# Visualization:
np.set_printoptions(precision=4, suppress=1)
correspondence = np.load('data/some_corresp.npz') # Loading correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1, im2 = plt.imread('data/im1.png'), plt.imread('data/im2.png')

F = eightpoint(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))
# epipolarMatchGUI(im1, im2, F) ## uncomment to visualize, comment out before submitting to gradescope
```

Output:
```
Optimization terminated successfully.
         Current function value: 0.000107
         Iterations: 8
         Function evaluations: 894
```

<!-- grade_id: q2_3_a, points: 3, grade: true, solution: false -->
```python
# Simple Tests to verify your implmentation:
correspondence = np.load('data/some_corresp.npz') # Loading correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
K1, K2 = intrinsics['K1'], intrinsics['K2']
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1, im2 = plt.imread('data/im1.png'), plt.imread('data/im2.png')
F = eightpoint(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))
x2, y2 = epipolarCorrespondence(im1, im2, F, 119, 217)

try:
    assert(np.linalg.norm(np.array([x2, y2]) - np.array([118, 181])) < 10)
    print('Test passed!')
except:
    raise AssertionError('Test Failed: error greater than 10!')

```

Output:
```
Optimization terminated successfully.
         Current function value: 0.000107
         Iterations: 8
         Function evaluations: 894
Test passed!
```

<!-- grade_id: q2_3_b, points: 7, grade: true, solution: false -->
```python
# Hidden Tests
```

### Q2.4 3D Visualization (3pt writeup, 7 pt implementation) 

Included in this homework  is a file data/templeCoords.npz which contains 288 hand-selected points from `im1` saved in the variables `x1` and `y1`.

Now, we can determine the 3D location of these point correspondences using the `triangulate` function. These 3D point locations can then plotted using the `Matplotlib` package (we have provided the necessary starter code for visualization). An example is shown here: 


|![alt](images/q3a.png) |![alt](images/q3b.png)|
|-|-|
|![alt](images/q3b.png) |![alt](images/q3c.png)|
|-|-|


<span style='color:red'>**Output:**</span> In your write-up: Take a few screenshots of the 3D visualization
so that the outline of the temple is clearly visible.

<!-- grade_id: cell-640b94de9f802a27, points: 0, grade: false, solution: true -->
```python
def compute3D_pts(pts1, intrinsics, F, im1, im2):
    '''
    Q2.4: Finding the 3D position of given points based on epipolar correspondence and triangulation
        Input:  pts1, chosen points from im1
                intrinsics, the intrinsics dictionary for calling epipolarCorrespondence
                F, the fundamental matrix
                im1, the first image
                im2, the second image
        Output: P (Nx3) the recovered 3D points
    
    ***
    Hints:
    (1) Use epipolarCorrespondence to find the corresponding point for [x1 y1] (find [x2, y2])
    (2) Now you have a set of corresponding points [x1, y1] and [x2, y2], you can compute the M2
        matrix and use triangulate to find the 3D points. 
    (3) Use the function findM2 to find the 3D points P (do not recalculate fundamental matrices)
    (4) As a reference, our solution's bet error is around ~2000 on the 3D points. 
    '''
    x1s_temple, y1s_temple = pts1[:, 0], pts1[:, 1]
    # ----- TODO -----
    # YOUR CODE HERE
    x2s_temple, y2s_temple = [], []
    for (x1,y1) in list(zip(x1s_temple, y1s_temple)):

        x2, y2 = epipolarCorrespondence(im1, im2, F, x1, y1)
        x2s_temple.append(x2)
        y2s_temple.append(y2)

    pts2 = np.array(list(zip(x2s_temple, y2s_temple)))
    M2, C2, P = find_M2(F, pts1, pts2, intrinsics)
    # raise NotImplementedError()
    return P


```

```python
# Visualization:
correspondence = np.load('data/some_corresp.npz') # Loading correspondences
intrinsics = np.load('data/intrinsics.npz') # Loading the intrinscis of the camera
templeCoords = np.load("data/templeCoords.npz")
pts1, pts2 = correspondence['pts1'], correspondence['pts2']
im1, im2 = plt.imread('data/im1.png'), plt.imread('data/im2.png')

F = eightpoint(pts1, pts2, M=np.max([*im1.shape, *im2.shape]))

pts1 = np.hstack([templeCoords["x1"], templeCoords["y1"]])
P = compute3D_pts(pts1, intrinsics, F, im1, im2)

# plot_3D(P) ## uncomment to visualize, comment out before submitting to gradescope

```

Output:
```
Optimization terminated successfully.
         Current function value: 0.000107
         Iterations: 8
         Function evaluations: 894
[[    0.     -1525.9       78.13       0.    ]
 [ 1520.4        0.      -111.68       0.    ]
 [  -27.4485 -1508.3883  -245.614   1532.9965]
 [-1519.9694    14.8087   121.2607   -48.8178]]
[[    0.     -1525.9       78.13       0.    ]
 [ 1520.4        0.      -111.68       0.    ]
 [  -27.4485 -1508.3883  -245.614  -1532.9965]
 [-1519.9694    14.8087   121.2607    48.8178]]
[[    0.     -1525.9       78.13       0.    ]
 [ 1520.4        0.      -111.68       0.    ]
 [   51.7203 -1496.056   -308.9604  1532.9965]
 [ 1519.1965    80.8669  -103.6004   -48.8178]]
[[    0.     -1525.9       78.13       0.    ]
 [ 1520.4        0.      -111.68       0.    ]
 [   51.7203 -1496.056   -308.9604 -1532.9965]
 [ 1519.1965    80.8669  -103.6004    48.8178]]
Best Error 530.871781133902
```

```python
plot_3D(P)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-53059fb55b60> in <module>
----> 1 plot_3D(P)

NameError: name 'plot_3D' is not defined
```

<!-- grade_id: q2_4_a, points: 7, grade: true, solution: false -->
```python
# Hidden Tests
```

### Q2.5 Extra Credits: More 3D Visualization (10pts writeup) 

Similar to the Q2.4 visualization of the Temple image, show the 3D reconstruction of any other object using your own images. Feel free to use the images and camera parameters from the [Middlebury multiview dataset](http://vision.middlebury.edu/mview/data/) or other public datasets. Please restrict to only using 2 images and atleast 10 correspondences in both the images.



<span style='color:red'>**Output:**</span> In your write-up: 
- Show the two images of the object from different views.
- Visualize the 2D correspondences on both the images.
- Take a few screenshots of the 3D visualization of the reconstruction.